<!DOCTYPE html>
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<html lang="en-US">

<head xmlns:og="http://ogp.me/ns#" xmlns:book="https://ogp.me/ns/book#">
  <meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
  <title>SB Parsons Exercises</title>
  <meta name="Keywords" content="Authored in PreTeXt">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <meta property="og:type" content="book">
  <meta property="book:title" content="PreTeXt Sample Book">
  <meta property="book:author" content="Thomas W. Judson">
  <script src="https://sagecell.sagemath.org/static/embedded_sagecell.js"></script>
  <script>var runestoneMathReady = new Promise((resolve) => window.rsMathReady = resolve);
    window.MathJax = {
      tex: {
        inlineMath: [['\\(', '\\)']],
        tags: "none",
        tagSide: "right",
        tagIndent: ".8em",
        packages: { '[+]': ['base', 'extpfeil', 'ams', 'amscd', 'color', 'newcommand', 'knowl'] }
      },
      options: {
        ignoreHtmlClass: "tex2jax_ignore|ignore-math",
        processHtmlClass: "process-math",
        renderActions: {
          findScript: [10, function (doc) {
            document.querySelectorAll('script[type^="math/tex"]').forEach(function (node) {
              var display = !!node.type.match(/; *mode=display/);
              var math = new doc.options.MathItem(node.textContent, doc.inputJax[0], display);
              var text = document.createTextNode('');
              node.parentNode.replaceChild(text, node);
              math.start = { node: text, delim: '', n: 0 };
              math.end = { node: text, delim: '', n: 0 };
              doc.math.push(math);
            });
          }, '']
        },
      },
      chtml: {
        scale: 0.98,
        mtextInheritFont: true
      },
      loader: {
        load: ['input/asciimath', '[tex]/extpfeil', '[tex]/amscd', '[tex]/color', '[tex]/newcommand', '[pretext]/mathjaxknowl3.js'],
        paths: { pretext: "https://pretextbook.org/js/lib" },
      },
      startup: {
        pageReady() {
          return MathJax.startup.defaultPageReady().then(function () {
            console.log("in ready function");
            rsMathReady();
          }
          )
        }
      },
    };
  </script>
  <script src="https://cdn.jsdelivr.net/npm/mathjax@3/es5/tex-chtml.js"></script>
  <link href="https://cdnjs.cloudflare.com/ajax/libs/prism/1.26.0/themes/prism.css" rel="stylesheet">
  <script src="https://cdnjs.cloudflare.com/ajax/libs/prism/1.26.0/components/prism-core.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/prism/1.26.0/plugins/autoloader/prism-autoloader.min.js"></script>
  <script src="https://unpkg.com/lunr/lunr.js"></script>
  <script src="lunr-pretext-search-index.js" async=""></script>
  <script src="https://pretextbook.org/js/0.2/pretext_search.js"></script>
  <link href="../pretext_css_core/pretext_search.css" rel="stylesheet" type="text/css">
  <script>js_version = 0.2</script>
  <script src="https://pretextbook.org/js/lib/jquery.min.js"></script>
  <script src="https://pretextbook.org/js/lib/jquery.sticky.js"></script>
  <script src="https://pretextbook.org/js/lib/jquery.espy.min.js"></script>
  <script src="https://pretextbook.org/js/0.2/pretext.js"></script>
  <script>miniversion = 0.1</script>
  <script src="https://pretextbook.org/js/0.2/pretext_add_on.js?x=1"></script>
  <script src="https://pretextbook.org/js/0.2/user_preferences.js"></script>
  <script src="https://pretextbook.org/js/lib/knowl.js"></script>
  <!--knowl.js code controls Sage Cells within knowls-->
  <script>sagecellEvalName = 'Evaluate (Sage)';
  </script>
  <link rel="preconnect" href="https://fonts.googleapis.com">
  <link rel="preconnect" href="https://fonts.gstatic.com" crossorigin="">
  <link
    href="https://fonts.googleapis.com/css2?family=Inconsolata:wght@400;700&amp;family=Noto+Serif:ital,wght@0,400;0,700;1,400;1,700&amp;family=Tinos:ital,wght@0,400;0,700;1,400;1,700&amp;display=swap"
    rel="stylesheet">
  <link href="https://fonts.cdnfonts.com/css/dejavu-serif" rel="stylesheet">
  <link
    href="https://fonts.googleapis.com/css2?family=Roboto+Serif:opsz,wdth,wght@8..144,50..150,100..900&amp;display=swap"
    rel="stylesheet">
  <link href="https://fonts.googleapis.com/css2?family=Open+Sans:wdth,wght@75..100,300..800&amp;display=swap"
    rel="stylesheet">
  <link href="../pretext_css_core/pretext.css" rel="stylesheet" type="text/css">
  <link href="../pretext_css_core/shell_default.css" rel="stylesheet" type="text/css">
  <link href="../pretext_css_core/banner_default.css" rel="stylesheet" type="text/css">
  <link href="../pretext_css_core/navbar_default.css" rel="stylesheet" type="text/css">
  <link href="../pretext_css_core/toc_default.css" rel="stylesheet" type="text/css">
  <link href="../pretext_css_core/knowls_default.css" rel="stylesheet" type="text/css">
  <link href="../pretext_css_core/style_default.css" rel="stylesheet" type="text/css">
  <link href="../pretext_css_core/colors_default.css" rel="stylesheet" type="text/css">
  <link href="../pretext_css_core/setcolors.css" rel="stylesheet" type="text/css">
  <!--** eBookCongig is necessary to configure interactive       **-->
  <!--** Runestone components to run locally in reader's browser **-->
  <!--** No external communication:                              **-->
  <!--**     log level is 0, Runestone Services are disabled     **-->
  <script type="text/javascript">
    eBookConfig = {};
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'http://127.0.0.1:8000';
    eBookConfig.course = 'PTX_Course_Title_Here';
    eBookConfig.basecourse = 'PTX_Base_Course';
    eBookConfig.isLoggedIn = false;
    eBookConfig.email = '';
    eBookConfig.isInstructor = false;
    eBookConfig.logLevel = 0;
    eBookConfig.username = '';
    eBookConfig.readings = null;
    eBookConfig.activities = null;
    eBookConfig.downloadsEnabled = false;
    eBookConfig.allow_pairs = false;
    eBookConfig.enableScratchAC = false;
    eBookConfig.build_info = "";
    eBookConfig.python3 = null;
    eBookConfig.runestone_version = '6.5.0';
    eBookConfig.jobehost = '';
    eBookConfig.proxyuri_runs = '';
    eBookConfig.proxyuri_files = '';
    eBookConfig.enable_chatcodes = false;
  </script>
  <!--*** Runestone Services ***-->
  <script type="text/javascript"
    src="https://runestone.academy/cdn/runestone/6.5.0/runtime.ead78978a4e3c14b.bundle.js"></script>
  <script type="text/javascript"
    src="https://runestone.academy/cdn/runestone/6.5.0/637.d54be67956c5c660.bundle.js"></script>
  <script type="text/javascript"
    src="https://runestone.academy/cdn/runestone/6.5.0/runestone.9ae535a5cdffdcb5.bundle.js"></script>
  <link rel="stylesheet" type="text/css" href="https://runestone.academy/cdn/runestone/6.5.0/637.0fa6cababf056764.css">
  <link rel="stylesheet" type="text/css"
    href="https://runestone.academy/cdn/runestone/6.5.0/runestone.828206f80e9e88d8.css">
    
  <link href="../pretext_css_core/pretext_add_on.css" rel="stylesheet" type="text/css">
</head>

<body id="PTXSB" class="pretext book ignore-math">
  <a class="assistive" href="#ptx-content">Skip to main content</a>
  <header id="ptx-masthead">
    <div class="ptx-banner">
      <a id="logo-link" href="http://abstract.pugetsound.edu" target="_blank"><img src="cover_aata_2014.png"
          alt="Logo image"></a>
      <div class="title-container">
        <h1 class="heading"><a href="sample-book.html"><span class="title">PreTeXt Sample Book:</span> <span
              class="subtitle">Abstract Algebra (SAMPLE ONLY)</span></a></h1>
        <p class="byline">Thomas W. Judson, Isaac Newton (Editor)</p>
      </div>
      <div id="searchresultsplaceholder" style="display: none">
        <button id="closesearchresults"
          onclick="document.getElementById('searchresultsplaceholder').style.display = 'none'; return false;">x</button>
        <h2>Search Results: <span id="searchterms"></span>
        </h2>
        <div id="searchempty"><span>No results.</span></div>
        <ol id="searchresults"></ol>
      </div>
    </div>
  </header>
  <nav id="ptx-navbar" class="navbar"><button class="toc-toggle button"
      aria-label="Show or hide table of contents"><span class="icon">☰</span><span
        class="name">Contents</span></button><a class="index-button button" href="index-1.html" title="Index"><span
        class="name">Index</span></a><button id="user-preferences-button" class="user-preferences-button button"
      title="Modify user preferences"><span id="theavatarbutton" class="name">You!</span>
      <div id="preferences_menu_holder" class="hidden">
        <ol id="preferences_menu" style="font-family: 'Roboto Serif', serif;">
          <li data-env="avatar" tabindex="-1">Choose avatar<div class="wrap_to_submenu"><span
                class="to_submenu">▻</span></div>
            <ol class="hidden avatar">
              <li data-val="You!" tabindex="-1">
                <span id="theYou!" class="avatarcheck">✔️</span>You!
              </li>
              <li data-val="😺" tabindex="-1">
                <span id="the😺" class="avatarcheck"></span>😺
              </li>
              <li data-val="👤" tabindex="-1">
                <span id="the👤" class="avatarcheck"></span>👤
              </li>
              <li data-val="👽" tabindex="-1">
                <span id="the👽" class="avatarcheck"></span>👽
              </li>
              <li data-val="🐶" tabindex="-1">
                <span id="the🐶" class="avatarcheck"></span>🐶
              </li>
              <li data-val="🐼" tabindex="-1">
                <span id="the🐼" class="avatarcheck"></span>🐼
              </li>
              <li data-val="🌈" tabindex="-1">
                <span id="the🌈" class="avatarcheck"></span>🌈
              </li>
            </ol>
          </li>
          <li data-env="fontfamily" tabindex="-1">Font family<div class="wrap_to_submenu"><span
                class="to_submenu">▻</span></div>
            <ol class="hidden fontfamily">
              <li data-val="face" data-change="OS" tabindex="-1" style="font-family: 'Open Sans'">
                <span id="theOS" class="ffcheck">✔️</span><span class="name">Open Sans</span><span class="sample">AaBbCc
                  123 PreTeXt</span>
              </li>
              <li data-val="face" data-change="RS" tabindex="-1" style="font-family: 'Roboto Serif'">
                <span id="theRS" class="ffcheck"></span><span class="name">Roboto Serif</span><span
                  class="sample">AaBbCc 123 PreTeXt</span>
              </li>
            </ol>
          </li>
          <li data-env="font" tabindex="-1">Adjust font<div class="wrap_to_submenu"><span class="to_submenu">▻</span>
            </div>
            <ol class="hidden fonts">
              <li>Size</li>
              <li><span id="thesize">12</span></li>
              <li data-val="size" data-change="-1" tabindex="-1" style="font-size: 80%">Smaller</li>
              <li data-val="size" data-change="1" tabindex="-1" style="font-size: 110%">Larger</li>
              <li>Width</li>
              <li><span id="thewdth">100</span></li>
              <li data-val="wdth" data-change="-5" tabindex="-1" style="font-variation-settings: 'wdth' 60">narrower
              </li>
              <li data-val="wdth" data-change="5" tabindex="-1" style="font-variation-settings: 'wdth' 150">wider</li>
              <li>Weight</li>
              <li><span id="thewght">400</span></li>
              <li data-val="wght" data-change="-50" tabindex="-1" style="font-weight: 200">thinner</li>
              <li data-val="wght" data-change="50" tabindex="-1" style="font-weight: 700">heavier</li>
              <li>Letter spacing</li>
              <li>
                <span id="thelspace">0</span><span class="byunits">/200</span>
              </li>
              <li data-val="lspace" data-change="-1" tabindex="-1">closer</li>
              <li data-val="lspace" data-change="1" tabindex="-1">f a r t h e r</li>
              <li>Word spacing</li>
              <li>
                <span id="thewspace">0</span><span class="byunits">/50</span>
              </li>
              <li data-val="wspace" data-change="-1" tabindex="-1">smaller gap </li>
              <li data-val="wspace" data-change="1" tabindex="-1">larger gap</li>
              <li>Line Spacing</li>
              <li>
                <span id="theheight">135</span><span class="byunits">/100</span>
              </li>
              <li data-val="height" data-change="-5" tabindex="-1" style="line-height: 1">closer<br>together</li>
              <li data-val="height" data-change="5" tabindex="-1" style="line-height: 1.75">further<br>apart</li>
            </ol>
          </li>
          <li data-env="atmosphere" tabindex="-1">Light/dark mode<div class="wrap_to_submenu"><span
                class="to_submenu">▻</span></div>
            <ol class="hidden atmosphere">
              <li data-val="default" tabindex="-1">
                <span id="thedefault" class="atmospherecheck">✔️</span>default
              </li>
              <li data-val="pastel" tabindex="-1">
                <span id="thepastel" class="atmospherecheck"></span>pastel
              </li>
              <li data-val="darktwilight" tabindex="-1">
                <span id="thedarktwilight" class="atmospherecheck"></span>twilight
              </li>
              <li data-val="dark" tabindex="-1">
                <span id="thedark" class="atmospherecheck"></span>dark
              </li>
              <li data-val="darkmidnight" tabindex="-1">
                <span id="thedarkmidnight" class="atmospherecheck"></span>midnight
              </li>
            </ol>
          </li>
          <li data-env="ruler" tabindex="-1">Reading ruler<div class="wrap_to_submenu"><span class="to_submenu">▻</span>
            </div>
            <ol class="hidden ruler">
              <li data-val="none" tabindex="-1">
                <span id="thenone" class="rulercheck">✔️</span>none
              </li>
              <li data-val="underline" tabindex="-1">
                <span id="theunderline" class="rulercheck"></span>underline
              </li>
              <li data-val="lunderline" tabindex="-1">
                <span id="thelunderline" class="rulercheck"></span>L-underline
              </li>
              <li data-val="greybar" tabindex="-1">
                <span id="thegreybar" class="rulercheck"></span>grey bar
              </li>
              <li data-val="lightbox" tabindex="-1">
                <span id="thelightbox" class="rulercheck"></span>light box
              </li>
              <li data-val="sunrise" tabindex="-1">
                <span id="thesunrise" class="rulercheck"></span>sunrise
              </li>
              <li data-val="sunriseunderline" tabindex="-1">
                <span id="thesunriseunderline" class="rulercheck"></span>sunrise underline
              </li>
              <li class="moveQ">Motion by:</li>
              <li data-val="mouse" tabindex="-1">
                <span id="themouse" class="motioncheck">✔️</span>follow the mouse
              </li>
              <li data-val="arrow" tabindex="-1">
                <span id="thearrow" class="motioncheck"></span>up/down arrows - not yet
              </li>
              <li data-val="eye" tabindex="-1">
                <span id="theeye" class="motioncheck"></span>eye tracking - not yet
              </li>
            </ol>
          </li>
        </ol>
      </div>
    </button><button
    id="searchbutton" type="button" class="button" onclick="doSearch()">🔍</button>
    <div class="searchbox">
      <div class="searchwidget">
        <input id="ptxsearch" type="text" name="terms" placeholder="Search" onchange="doSearch()">
      </div>
    </div>
    <div class="flex-spacer"></div>
    <div class="treebuttons"><a class="previous-button button" href="multiple-choice-exercises.html"
        title="Previous"><span class="icon">&lt;</span><span class="name">Prev</span></a><a class="up-button button"
        href="rune.html" title="Up"><span class="icon">^</span><span class="name">Up</span></a><a
        class="next-button button" href="horizontal-parsons-exercises.html" title="Next"><span
          class="name">Next</span><span class="icon">&gt;</span></a></div>
  </nav>
  <div id="latex-macros" class="hidden-content process-math" style="display:none"><span
      class="process-math">\(\newcommand{\identity}{\mathrm{id}}
      \newcommand{\notdivide}{{\not{\mid}}}
      \newcommand{\notsubset}{\not\subset}
      \newcommand{\lcm}{\operatorname{lcm}}
      \newcommand{\gf}{\operatorname{GF}}
      \newcommand{\inn}{\operatorname{Inn}}
      \newcommand{\aut}{\operatorname{Aut}}
      \newcommand{\Hom}{\operatorname{Hom}}
      \newcommand{\cis}{\operatorname{cis}}
      \newcommand{\chr}{\operatorname{char}}
      \newcommand{\Null}{\operatorname{Null}}
      \newcommand{\lt}{&lt;}
      \newcommand{\gt}{&gt;}
      \newcommand{\amp}{&amp;}
      \definecolor{fillinmathshade}{gray}{0.9}
      \newcommand{\fillinmath}[1]{\mathchoice{\colorbox{fillinmathshade}{$\displaystyle
      \phantom{\,#1\,}$}}{\colorbox{fillinmathshade}{$\textstyle
      \phantom{\,#1\,}$}}{\colorbox{fillinmathshade}{$\scriptstyle
      \phantom{\,#1\,}$}}{\colorbox{fillinmathshade}{$\scriptscriptstyle\phantom{\,#1\,}$}}}
      \newcommand{\sfrac}[2]{{#1}/{#2}}
      \)</span></div>
  <div class="ptx-page">
    <div id="ptx-sidebar" class="hidden">
      <nav id="ptx-toc" class="depth3 parts">
        <ul class="structural">
          <li>
            <div class="toc-item"><a href="frontmatter.html" class="internal"><span class="title">Front
                  Matter</span></a></div>
            <ul class="structural">
              <li>
                <div class="toc-item"><a href="front-colophon.html" class="internal"><span
                      class="title">Colophon</span></a></div>
              </li>
              <li>
                <div class="toc-item"><a href="author-bio-TWJ.html" class="internal"><span class="title">Author
                      Biography</span></a></div>
              </li>
              <li>
                <div class="toc-item"><a href="author-bio-RAB.html" class="internal"><span class="title">About Robert A.
                      Beezer</span></a></div>
              </li>
              <li>
                <div class="toc-item"><a href="dedication.html" class="internal"><span
                      class="title">Dedication</span></a></div>
              </li>
              <li>
                <div class="toc-item"><a href="acknowledgement.html" class="internal"><span
                      class="title">Acknowledgements</span></a></div>
              </li>
              <li>
                <div class="toc-item"><a href="preface.html" class="internal"><span class="title">Preface</span></a>
                </div>
              </li>
              <li>
                <div class="toc-item"><a href="contributors.html" class="internal"><span class="title">Contributors to
                      the 4<span class="process-math">\(^\mathrm{th}\)</span> Edition</span></a></div>
                <ul class="structural">
                  <li>
                    <div class="toc-item"><a href="contributors.html#contributors-1" class="internal"><span
                          class="title"></span></a></div>
                  </li>
                </ul>
              </li>
            </ul>
          </li>
          <li>
            <div class="toc-item"><a href="part-basics.html" class="internal"><span class="codenumber">I</span> <span
                  class="title">Basics</span></a></div>
            <ul class="structural">
              <li>
                <div class="toc-item"><a href="sets.html" class="internal"><span class="codenumber">1</span> <span
                      class="title">Preliminaries</span></a></div>
                <ul class="structural">
                  <li>
                    <div class="toc-item"><a href="section-note-on-proofs.html" class="internal"><span
                          class="codenumber">1.1</span> <span class="title">A Short Note on Proofs</span></a></div>
                    <ul class="structural">
                      <li>
                        <div class="toc-item"><a href="section-note-on-proofs.html#subsection-1" class="internal"><span
                              class="codenumber">1.1.1</span> <span class="title">Some Cautions and
                              Suggestions</span></a></div>
                      </li>
                    </ul>
                  </li>
                  <li>
                    <div class="toc-item"><a href="section-sets-and-equivalence-relations.html" class="internal"><span
                          class="codenumber">1.2</span> <span class="title">Sets and Equivalence Relations</span></a>
                    </div>
                    <ul class="structural">
                      <li>
                        <div class="toc-item"><a href="section-sets-and-equivalence-relations.html#subsection-2"
                            class="internal"><span class="codenumber">1.2.1</span> <span class="title">Set
                              Theory</span></a></div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="section-sets-and-equivalence-relations.html#subsection-3"
                            class="internal"><span class="codenumber">1.2.2</span> <span class="title">Cartesian
                              Products and Mappings</span></a></div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="section-sets-and-equivalence-relations.html#subsection-4"
                            class="internal"><span class="codenumber">1.2.3</span> <span class="title">Equivalence
                              Relations and Partitions</span></a></div>
                      </li>
                    </ul>
                  </li>
                  <li>
                    <div class="toc-item"><a href="sets-sage.html" class="internal"><span class="codenumber">1.3</span>
                        <span class="title">Sage</span></a></div>
                    <ul class="structural">
                      <li>
                        <div class="toc-item"><a href="sets-sage.html#subsection-5" class="internal"><span
                              class="codenumber">1.3.1</span> <span class="title">Executing Sage Commands</span></a>
                        </div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="sets-sage.html#subsection-6" class="internal"><span
                              class="codenumber">1.3.2</span> <span class="title">Immediate Help</span></a></div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="sets-sage.html#subsection-7" class="internal"><span
                              class="codenumber">1.3.3</span> <span class="title">Annotating Your Work</span></a></div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="sets-sage.html#subsection-8" class="internal"><span
                              class="codenumber">1.3.4</span> <span class="title">Lists</span></a></div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="sets-sage.html#subsection-9" class="internal"><span
                              class="codenumber">1.3.5</span> <span class="title">Lists of Integers</span></a></div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="sets-sage.html#subsection-10" class="internal"><span
                              class="codenumber">1.3.6</span> <span class="title">Saving and Sharing Your
                              Work</span></a></div>
                      </li>
                    </ul>
                  </li>
                  <li>
                    <div class="toc-item"><a href="exercises-sets.html" class="internal"><span
                          class="codenumber">1.4</span> <span class="title">Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="sets-sage-exercises.html" class="internal"><span
                          class="codenumber">1.5</span> <span class="title">Sage Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="references-1.html" class="internal"><span
                          class="codenumber">1.6</span> <span class="title">References and Suggested Readings</span></a>
                    </div>
                  </li>
                </ul>
              </li>
              <li>
                <div class="toc-item"><a href="integers.html" class="internal"><span class="codenumber">2</span> <span
                      class="title">The Integers</span></a></div>
                <ul class="structural">
                  <li>
                    <div class="toc-item"><a href="section-math-induction.html" class="internal"><span
                          class="codenumber">2.1</span> <span class="title">Mathematical Induction and Math in a Title
                          <span class="process-math">\(A\notsubset B\)</span></span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="section-division-algorithm.html" class="internal"><span
                          class="codenumber">2.2</span> <span class="title">The Division Algorithm</span></a></div>
                    <ul class="structural">
                      <li>
                        <div class="toc-item"><a href="section-division-algorithm.html#subsection-11"
                            class="internal"><span class="codenumber">2.2.1</span> <span class="title">The Euclidean
                              Algorithm</span></a></div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="section-division-algorithm.html#subsection-12"
                            class="internal"><span class="codenumber">2.2.2</span> <span class="title">Prime
                              Numbers</span></a></div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="section-division-algorithm.html#subsection-13"
                            class="internal"><span class="codenumber">2.2.3</span> <span class="title">Historical
                              Note</span></a></div>
                      </li>
                    </ul>
                  </li>
                  <li>
                    <div class="toc-item"><a href="integers-sage.html" class="internal"><span
                          class="codenumber">2.3</span> <span class="title">Sage</span></a></div>
                    <ul class="structural">
                      <li>
                        <div class="toc-item"><a href="integers-sage.html#subsection-14" class="internal"><span
                              class="codenumber">2.3.1</span> <span class="title">Division Algorithm</span></a></div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="integers-sage.html#subsection-15" class="internal"><span
                              class="codenumber">2.3.2</span> <span class="title">Greatest Common Divisor</span></a>
                        </div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="integers-sage.html#subsection-16" class="internal"><span
                              class="codenumber">2.3.3</span> <span class="title">Primes and Factoring</span></a></div>
                      </li>
                    </ul>
                  </li>
                  <li>
                    <div class="toc-item"><a href="exercises-integers.html" class="internal"><span
                          class="codenumber">2.4</span> <span class="title">Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="exercises-4.html" class="internal"><span
                          class="codenumber">2.5</span> <span class="title">Programming Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="integers-sage-exercises.html" class="internal"><span
                          class="codenumber">2.6</span> <span class="title">Sage Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="references-2.html" class="internal"><span
                          class="codenumber">2.7</span> <span class="title">References and Suggested Readings</span></a>
                    </div>
                  </li>
                </ul>
              </li>
            </ul>
          </li>
          <li>
            <div class="toc-item"><a href="part-algebra.html" class="internal"><span class="codenumber">II</span> <span
                  class="title">Algebra</span></a></div>
            <ul class="structural">
              <li>
                <div class="toc-item"><a href="groups.html" class="internal"><span class="codenumber">3</span> <span
                      class="title">Groups</span></a></div>
                <ul class="structural">
                  <li>
                    <div class="toc-item"><a href="section-mod-n-sym.html" class="internal"><span
                          class="codenumber">3.1</span> <span class="title">Integer Equivalence Classes and
                          Symmetries</span></a></div>
                    <ul class="structural">
                      <li>
                        <div class="toc-item"><a href="section-mod-n-sym.html#subsection-17" class="internal"><span
                              class="codenumber">3.1.1</span> <span class="title">The Integers mod <span
                                class="process-math">\(n\)</span></span></a></div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="section-mod-n-sym.html#subsection-18" class="internal"><span
                              class="codenumber">3.1.2</span> <span class="title">Symmetries</span></a></div>
                      </li>
                    </ul>
                  </li>
                  <li>
                    <div class="toc-item"><a href="section-groups-define.html" class="internal"><span
                          class="codenumber">3.2</span> <span class="title">Definitions and Examples</span></a></div>
                    <ul class="structural">
                      <li>
                        <div class="toc-item"><a href="section-groups-define.html#subsection-19" class="internal"><span
                              class="codenumber">3.2.1</span> <span class="title">Definition of a Group</span></a></div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="section-groups-define.html#subsection-20" class="internal"><span
                              class="codenumber">3.2.2</span> <span class="title">Examples of Groups</span></a></div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="section-groups-define.html#subsection-21" class="internal"><span
                              class="codenumber">3.2.3</span> <span class="title">Basic Properties of Groups</span></a>
                        </div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="section-groups-define.html#subsection-22" class="internal"><span
                              class="codenumber">3.2.4</span> <span class="title">Historical Note</span></a></div>
                      </li>
                    </ul>
                  </li>
                  <li>
                    <div class="toc-item"><a href="section-subgroups.html" class="internal"><span
                          class="codenumber">3.3</span> <span class="title">Subgroups</span></a></div>
                    <ul class="structural">
                      <li>
                        <div class="toc-item"><a href="section-subgroups.html#subsection-23" class="internal"><span
                              class="codenumber">3.3.1</span> <span class="title">Definitions and Examples</span></a>
                        </div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="section-subgroups.html#subsection-24" class="internal"><span
                              class="codenumber">3.3.2</span> <span class="title">Some Subgroup Theorems</span></a>
                        </div>
                      </li>
                    </ul>
                  </li>
                  <li>
                    <div class="toc-item"><a href="groups-sage.html" class="internal"><span
                          class="codenumber">3.4</span> <span class="title">Sage</span></a></div>
                    <ul class="structural">
                      <li>
                        <div class="toc-item"><a href="groups-sage.html#subsection-25" class="internal"><span
                              class="codenumber">3.4.1</span> <span class="title">Integers mod n</span></a></div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="groups-sage.html#subsection-26" class="internal"><span
                              class="codenumber">3.4.2</span> <span class="title">Groups of symmetries</span></a></div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="groups-sage.html#subsection-27" class="internal"><span
                              class="codenumber">3.4.3</span> <span class="title">Quaternions</span></a></div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="groups-sage.html#subsection-28" class="internal"><span
                              class="codenumber">3.4.4</span> <span class="title">Subgroups</span></a></div>
                      </li>
                    </ul>
                  </li>
                  <li>
                    <div class="toc-item"><a href="exercises-groups.html" class="internal"><span
                          class="codenumber">3.5</span> <span class="title">Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="exercises-7.html" class="internal"><span
                          class="codenumber">3.6</span> <span class="title">Additional Exercises: Detecting
                          Errors</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="groups-sage-exercises.html" class="internal"><span
                          class="codenumber">3.7</span> <span class="title">Sage Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="references-3.html" class="internal"><span
                          class="codenumber">3.8</span> <span class="title">References and Suggested Readings</span></a>
                    </div>
                  </li>
                </ul>
              </li>
              <li>
                <div class="toc-item"><a href="cyclic.html" class="internal"><span class="codenumber">4</span> <span
                      class="title">Cyclicity</span></a></div>
                <ul class="structural">
                  <li>
                    <div class="toc-item"><a href="section-11.html" class="internal"><span class="codenumber">4.1</span>
                        <span class="title">Cyclic groups</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="section-12.html" class="internal"><span class="codenumber">4.2</span>
                        <span class="title">Subgroups of a Cyclic Group</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="section-13.html" class="internal"><span class="codenumber">4.3</span>
                        <span class="title">Cyclic Groups of Complex Numbers</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="section-14.html" class="internal"><span class="codenumber">4.4</span>
                        <span class="title">Large Powers of Integers</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="exercises-cyclic.html" class="internal"><span
                          class="codenumber">4.5</span> <span class="title">Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="exercises-10.html" class="internal"><span
                          class="codenumber">4.6</span> <span class="title">Programming Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="cyclic-sage-exercises.html" class="internal"><span
                          class="codenumber">4.7</span> <span class="title">Sage Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="references-4.html" class="internal"><span
                          class="codenumber">4.8</span> <span class="title">References and Suggested Readings</span></a>
                    </div>
                  </li>
                </ul>
              </li>
              <li>
                <div class="toc-item"><a href="rune.html" class="internal"><span class="codenumber">5</span> <span
                      class="title">Runestone Testing</span></a></div>
                <ul class="structural">
                  <li>
                    <div class="toc-item"><a href="activecode.html" class="internal"><span class="codenumber">5.1</span>
                        <span class="title">ActiveCode</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="code-lens.html" class="internal"><span class="codenumber">5.2</span>
                        <span class="title">Code Lens</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="coding-exercises.html" class="internal"><span
                          class="codenumber">5.3</span> <span class="title">Coding Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="data-files.html" class="internal"><span class="codenumber">5.4</span>
                        <span class="title">Data Files</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="videos-in-runestone.html" class="internal"><span
                          class="codenumber">5.5</span> <span class="title">YouTube Videos</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="rune-nested-divisions.html" class="internal"><span
                          class="codenumber">5.6</span> <span class="title">Deeper</span></a></div>
                    <ul class="structural">
                      <li>
                        <div class="toc-item"><a href="rune-nested-divisions.html#subsection-29" class="internal"><span
                              class="codenumber">5.6.1</span> <span class="title">Subsection One</span></a></div>
                      </li>
                      <li>
                        <div class="toc-item"><a href="rune-nested-divisions.html#subsection-30" class="internal"><span
                              class="codenumber">5.6.2</span> <span class="title">Subsection Two</span></a></div>
                      </li>
                    </ul>
                  </li>
                  <li>
                    <div class="toc-item"><a href="true-false-exercises.html" class="internal"><span
                          class="codenumber">5.7</span> <span class="title">True/False Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="multiple-choice-exercises.html" class="internal"><span
                          class="codenumber">5.8</span> <span class="title">Multiple Choice Exercises</span></a></div>
                  </li>
                  <li class="active">
                    <div class="toc-item"><a href="parsons-exercises.html" class="internal"><span
                          class="codenumber">5.9</span> <span class="title">Parsons Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="horizontal-parsons-exercises.html" class="internal"><span
                          class="codenumber">5.10</span> <span class="title">Horizontal Parsons Exercises</span></a>
                    </div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="matching-exercises.html" class="internal"><span
                          class="codenumber">5.11</span> <span class="title">Matching Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="clickable-area-exercises.html" class="internal"><span
                          class="codenumber">5.12</span> <span class="title">Clickable Area Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="short-answer-exercises.html" class="internal"><span
                          class="codenumber">5.13</span> <span class="title">Short Answer Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="fill-in--exercises.html" class="internal"><span
                          class="codenumber">5.14</span> <span class="title">Fill-In Exercises</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="exercises-timed.html" class="internal"><span
                          class="codenumber">5.15</span> <span class="title">Exercises that are Timed</span></a></div>
                    <ul class="structural">
                      <li>
                        <div class="toc-item"><a href="exercises-timed.html#timed-exercises" class="internal"><span
                              class="codenumber">5.15</span> <span class="title">Timed Exercises</span></a></div>
                      </li>
                    </ul>
                  </li>
                  <li>
                    <div class="toc-item"><a href="interactive-projects.html" class="internal"><span
                          class="codenumber">5.16</span> <span class="title">Projects and Friends</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="section-23.html" class="internal"><span
                          class="codenumber">5.17</span> <span class="title">Expedited Samples</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="reading-questions-1.html" class="internal"><span
                          class="codenumber">5.18</span> <span class="title">Reading Questions</span></a></div>
                  </li>
                  <li>
                    <div class="toc-item"><a href="atomic-video.html" class="internal"><span
                          class="codenumber">5.19</span> <span class="title">YouTube Video Embedding</span></a></div>
                  </li>
                </ul>
              </li>
            </ul>
          </li>
          <li>
            <div class="toc-item"><a href="backmatter.html" class="internal"><span class="title">Back Matter</span></a>
            </div>
            <ul class="structural">
              <li>
                <div class="toc-item"><a href="appendix-1.html" class="internal"><span class="codenumber">A</span> <span
                      class="title">Notation</span></a></div>
              </li>
              <li>
                <div class="toc-item"><a href="solutions-1.html" class="internal"><span class="codenumber">B</span>
                    <span class="title">Hints and Answers to Selected Odd Exercises</span></a></div>
              </li>
              <li>
                <div class="toc-item"><a href="solutions-2.html" class="internal"><span class="codenumber">C</span>
                    <span class="title">Hints and Answers to Selected Even Exercises</span></a></div>
              </li>
              <li>
                <div class="toc-item"><a href="solutions-3.html" class="internal"><span class="codenumber">D</span>
                    <span class="title">Hints and Answers to Selected Exercises</span></a></div>
              </li>
              <li>
                <div class="toc-item"><a href="appendix-2.html" class="internal"><span class="codenumber">E</span> <span
                      class="title">A Structured Appendix</span></a></div>
                <ul class="structural">
                  <li>
                    <div class="toc-item"><a href="section-25.html" class="internal"><span class="codenumber">E.1</span>
                        <span class="title">A Section in an Appendix</span></a></div>
                    <ul class="structural">
                      <li>
                        <div class="toc-item"><a href="section-25.html#subsection-31" class="internal"><span
                              class="codenumber">E.1.1</span> <span class="title">A Subsection in a Section in an
                              Appendix</span></a></div>
                        <ul class="structural">
                          <li>
                            <div class="toc-item"><a href="section-25.html#subsubsection-1" class="internal"><span
                                  class="codenumber">E.1.1.1</span> <span class="title">A Subsubsection in a Subsection
                                  in a Section in an Appendix</span></a></div>
                          </li>
                        </ul>
                      </li>
                    </ul>
                  </li>
                  <li>
                    <div class="toc-item"><a href="section-26.html" class="internal"><span class="codenumber">E.2</span>
                        <span class="title">Numbering in the Back Matter</span></a></div>
                  </li>
                </ul>
              </li>
              <li>
                <div class="toc-item"><a href="appendix-gfdl.html" class="internal"><span class="codenumber">F</span>
                    <span class="title">GNU Free Documentation License</span></a></div>
              </li>
              <li>
                <div class="toc-item"><a href="index-1.html" class="internal"><span class="title">Index</span></a></div>
              </li>
              <li>
                <div class="toc-item"><a href="back-colophon.html" class="internal"><span
                      class="title">Colophon</span></a></div>
              </li>
            </ul>
          </li>
        </ul>
      </nav>
    </div>
    <main class="ptx-main">
      <div id="ptx-content" class="ptx-content">
        <section class="section" id="activecode">
          <h2 class="heading hide-type">
            <span class="type">Section</span> <span class="codenumber">5.1</span> <span class="title">ActiveCode</span>
          </h2>

          <div class="para" id="p-800">A Java program will only be interactive if hosted on a Runestone server.</div>
          <figure class="listing figure-like" id="program-activecode-java">
            <pre class="program"><code class="language-java">import javax.swing.JFrame;  //Importing class JFrame
import javax.swing.JLabel;  //Importing class JLabel
public class HelloWorld {
    public static void main(String[] args) {

        JFrame frame = new JFrame();           //Creating frame
        frame.setTitle("Hi!");                 //Setting title frame
        frame.add(new JLabel("Hello, world!"));//Adding text to frame
        frame.pack();                          //Setting size to smallest
        frame.setLocationRelativeTo(null);     //Centering frame
        //let's make this look more like real Java line length...
        app.getFrameFactory().buildFrame(app.getConfig().getDimensions()).setVisible(true);  //Showing frame
        app.getFrameFactory().buildFrame(app.getConfig().getDimensions()).setVisible(true);  //Showing frame
        app.getFrameFactory().buildFrame(app.getConfig().getDimensions()).setVisible(true);  //Showing frame
        app.getFrameFactory().buildFrame(app.getConfig().getDimensions()).setVisible(true);  //Showing frame
    }
}
</code></pre>
            <figcaption><span class="type">Listing</span><span class="space"> </span><span class="codenumber">5.1.3<span
                  class="period">.</span></span><span class="space"> </span>A Java program, interactive on a <span
                class="booktitle">Runestone</span> server</figcaption>
          </figure>




          <div class="para" id="p-805">The following Python program is in a <code
              class="code-inline tex2jax_ignore">&lt;listing&gt;</code> since we will want to reference it shortly. The
            program does not do very much, it just defines four variables whose values are lists of statistics. It
            should run, and there will be no syntax errors, but it is a bit boring since there is no output.</div>
          <figure class="listing figure-like" id="listing-python-included">
            <div class="ptx-runestone-container">
              <div class="runestone explainer ac_section ">
                <div data-component="activecode" id="statistics"><textarea data-lang="python" data-timelimit="25000"
                    data-audio="" data-coach="true" style="visibility: hidden;" id="statistics_editor"
                    data-question_label="" data-codelens="true">loan_amount = [1250.0, 500.0, 1450.0, 200.0, 700.0, 100.0, 250.0, 225.0, 1200.0, 150.0, 600.0, 300.0, 700.0, 125.0, 650.0, 175.0, 1800.0, 1525.0, 575.0, 700.0, 1450.0, 400.0, 200.0, 1000.0, 350.0]

country_name = ['Azerbaijan', 'El Salvador', 'Bolivia', 'Paraguay', 'El Salvador', 'Philippines', 'Philippines', 'Nicaragua', 'Guatemala', 'Philippines', 'Paraguay', 'Philippines', 'Bolivia', 'Philippines', 'Philippines', 'Madagascar', 'Georgia', 'Uganda', 'Kenya', 'Tajikistan', 'Jordan', 'Kenya', 'Philippines', 'Ecuador', 'Kenya']

time_to_raise = [193075.0, 1157108.0, 1552939.0, 244945.0, 238797.0, 1248909.0, 773599.0, 116181.0, 2288095.0, 51668.0, 26717.0, 48030.0, 1839190.0, 71117.0, 580401.0, 800427.0, 1156218.0, 1166045.0, 2924705.0, 470622.0, 24078.0, 260044.0, 445938.0, 201408.0, 2370450.0]

num_lenders_total = [38, 18, 51, 3, 21, 1, 10, 8, 42, 1, 18, 6, 28, 5, 16, 7, 54, 1, 18, 22, 36, 12, 8, 24, 8]
</textarea></div>
              </div>
            </div>
            <figcaption><span class="type">Listing</span><span class="space"> </span><span class="codenumber">5.1.8<span
                  class="period">.</span></span><span class="space"> </span>A Python program that defines some
              statistics</figcaption>
          </figure>

          <article class="exercise exercise-like" id="exercise-python-including"><a href="" data-knowl=""
              class="id-ref exercise-knowl original" data-refid="hk-exercise-python-including">
              <h3 class="heading">
                <span class="type">Checkpoint</span><span class="space"> </span><span
                  class="codenumber">5.1.9</span><span class="period">.</span><span class="space"> </span><span
                  class="title">A Python program, including another.</span>
              </h3>
            </a></article>
          <div class="hidden-content tex2jax_ignore" id="hk-exercise-python-including">
            <article class="exercise exercise-like">
              <div class="ptx-runestone-container">
                <div class="runestone explainer ac_section ">
                  <div data-component="activecode" id="exercise-python-including">
                    <div class="ac_question col-md-12" id="exercise-python-including_question">
                      <div class="para" id="p-808">Compute the total amount of money loaned and store it in the variable
                        <code class="code-inline tex2jax_ignore">loan_total</code>.</div>
                    </div>
                    <textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true"
                      style="visibility: hidden;" id="exercise-python-including_editor" data-question_label=""
                      data-codelens="true" data-include="statistics">loan_total = 0
for loan in loan_amount:
    loan_total += loan
print(loan_total)
====
from unittest.gui import TestCaseGui

class MyTests(TestCaseGui):

    def testOne(self):
        self.assertTrue('loan_total' in self.getEditorText(), "you need a loan_total variable")
        self.assertEqual(loan_total, sum(loan_amount), "Use the accumulator pattern to add up all the loans")
        self.assertFalse('sum(' in self.getEditorText(), "you may not use sum()")

MyTests().main()
</textarea>
                  </div>
                </div>
              </div>
            </article>
          </div>

          <div class="para" id="p-811">A nonsense paragraph just to check on spacing. A nonsense paragraph just to check
            on spacing. A nonsense paragraph just to check on spacing. A nonsense paragraph just to check on spacing.
          </div>
        </section>

        <section class="exercises" id="parsons-exercises">
          <h2 class="heading hide-type">
            <span class="type">Exercises</span> <span class="codenumber">5.9</span> <span class="title">Parsons
              Exercises</span>
          </h2>
          <article class="exercise exercise-like" id="number-theory-proof">
            <h3 class="heading">
              <span class="codenumber">1<span class="period">.</span></span><span class="space"> </span><span
                class="title">Parsons Problem, Mathematical Proof.</span>
            </h3>
            <div class="para" id="p-1997"><div class="autopermalink" onclick="copyPermalink(this)" data-description="Paragraph"><a href="https://pretextbook.org/doc/guide/html/extra-stylesheet.html#p-1997" title="Copy permalink for Paragraph">🔗</a></div>There are situations when an “extra” <abbr class="initialism">XSL</abbr> stylesheet is necessary for processing your <abbr class="initialism">XML</abbr> source.  A good example is the supported method for styling <abbr class="initialism">PDF</abbr> output, described at <a href="latex-styles.html" class="internal" title="Chapter 41: LaTeX Styles">Chapter&nbsp;41</a>.  Historically, various customizations have been supported by “thin” XSL stylesheets (<a href="publisher-extra-stylesheet.html" class="internal" title="Section 28.2: Extra Stylesheets">Section&nbsp;28.2</a>), but as certain themes have become apparent we have moved these customizations to simpler techniques (e.g. publication files, <a href="publication-file.html" class="internal" title="Section 26.1: Publication File">Section&nbsp;26.1</a>).  Other than styling <abbr class="initialism">PDF</abbr>, you should think carefully about if an additional stylesheet is necessary.</div>
            <div class="ptx-runestone-container">
              <div class="runestone parsons_section">
                <div data-component="parsons" class="parsons" id="number-theory-proof">
                  <div class="parsons_question parsons-text">
                    <div class="para" id="p-900">Create a proof of the theorem: If <span
                        class="process-math">\(n\)</span> is an even number, then <span class="process-math">\(n\equiv
                        0\mod 2\text{.}\)</span>
                    </div>
                  </div>
                  <pre class="parsonsblocks" data-question_label="" style="visibility: hidden;" data-language="natural"
                    data-adaptive="true" data-noindent="true"><div class="para" id="p-901">Suppose <span class="process-math">\(n\)</span> is even.</div>
---
<div class="para" id="p-903">Then there exists an <span class="process-math">\(m\)</span> so that <span class="process-math">\(n = 2m\text{.}\)</span>
</div>
---
<div class="para" id="p-902">Then <span class="process-math">\(n\)</span> is a prime number.</div> #paired
---
<div class="para" id="p-904">Then there exists an <span class="process-math">\(m\)</span> so that <span class="process-math">\(n = 2m + 1\text{.}\)</span>
</div> #paired
---
<div class="para" id="p-905">Click the heels of your ruby slippers together three times.</div> #distractor
---
<div class="para" id="p-906">So <span class="process-math">\(n = 2m + 0\text{.}\)</span>
</div> <div class="para" id="p-907">This is a superfluous second paragraph in this block.</div>
---
<div class="para" id="p-908">Thus <span class="process-math">\(n\equiv 0\mod 2\text{.}\)</span>
</div>
---
<div class="para logical" id="p-909">
<div class="para">And a little bit of irrelevant multi-line math</div>
<div class="displaymath process-math" id="md-45">
\begin{align*}
c^2&amp;a^2+b^2\\
&amp;x^2+y^2\text{.}
\end{align*}
</div>
</div> #distractor</pre>
                </div>
              </div>
            </div>
            <div class="solutions">
              <a href="" data-knowl="" class="id-ref hint-knowl original" data-refid="hk-hint-40" id="hint-40"><span
                  class="type">Hint</span><span class="period">.</span></a>
              <div class="hidden-content tex2jax_ignore" id="hk-hint-40">
                <div class="hint solution-like">Dorothy will not be much help with this proof.</div>
              </div>
            </div>
          </article>
          <article class="exercise exercise-like" id="prime-number-program-indent-yes">
            <h3 class="heading">
              <span class="codenumber">2<span class="period">.</span></span><span class="space"> </span><span
                class="title">Parsons Problem, Programming.</span>
            </h3>
            <div class="ptx-runestone-container">
              <div class="runestone parsons_section">
                <div data-component="parsons" class="parsons" id="prime-number-program-indent-yes">
                  <div class="parsons_question parsons-text">
                    <div class="para" id="p-910">The Sieve of Eratosthenes computes prime numbers by starting with a
                      finite list of the integers bigger than 1. The first member of the list is a prime and is
                      saved/recorded. Then all multiples of that prime (which not a prime, excepting the prime itself!)
                      are removed from the list. Now the first number remaining in the list is the next prime number.
                      And the process repeats.</div>
                    <div class="para" id="p-911">The code blocks below can be rearranged to form one of the many
                      possible programs to implement this algorithm to compute a list of all the primes less than <span
                        class="process-math">\(250\text{.}\)</span> [Ed. this version of this problem requires the
                      reader to provide the necessary indentation.]</div>
                    <div class="para" id="p-912">This reprises <a href="" class="xref"
                        data-knowl="./knowl/sieve-primes.html"
                        title="Exercise 2.5.1: The Sieve of Eratosthenes">Exercise 2.5.1</a>.</div>
                  </div>
                  <pre class="parsonsblocks" data-question_label="" style="visibility: hidden;" data-language="python"
                    data-adaptive="true" data-noindent="false">n = 250
---
primes = []
candidates = list(range(2,n))
---
candidates = []
primes = list(range(2,n)) #paired
---
primes = candidates + [p] #distractor
---
while candidates:
---
    p = candidates[0]
    primes.append(p)
---
    for nonprime in range(p, n, p):
---
        if nonprime in candidates:
            candidates.remove(nonprime) #let's test a really, really, really, really long line
---
print(primes)</pre>
                </div>
              </div>
            </div>
          </article>
        </section>
      </div>

      <div class="ptx-content-footer">
        <a class="previous-button button" href="acknowledgement.html" title="Previous"><span class="icon">&lt;</span><span class="name">Prev</span></a><a class="top-button button" href="#" title="Top"><span class="icon">^</span><span class="name">Top</span></a><a class="next-button button" href="contributors.html" title="Next"><span class="name">Next</span><span class="icon">&gt;</span></a>
        </div>

    </main>
  </div>
  <div id="ptx-page-footer">
    <div class="ptx-page-footer-contents">
    <a class="pretext-link" href="https://pretextbook.org" title="PreTeXt">
      <div class="logo"><svg xmlns="http://www.w3.org/2000/svg" width="100%" height="100%" viewBox="338 3000 8772 6866">
          <g style="stroke-width:.025in; stroke:black; fill:none">
            <polyline points="472,3590 472,9732 "
              style="stroke:#000000;stroke-width:174; stroke-linejoin:miter; stroke-linecap:round; "></polyline>
            <path style="stroke:#000000;stroke-width:126;stroke-linecap:butt;"
              d="M 4724,9448 A 4660 4660  0  0  1  8598  9259"></path>
            <path style="stroke:#000000;stroke-width:174;stroke-linecap:butt;"
              d="M 4488,9685 A 4228 4228  0  0  0   472  9732"></path>
            <path style="stroke:#000000;stroke-width:126;stroke-linecap:butt;"
              d="M 4724,3590 A 4241 4241  0  0  1  8598  3496"></path>
            <path style="stroke:#000000;stroke-width:126;stroke-linecap:round;"
              d="M 850,3496  A 4241 4241  0  0  1  4724  3590"></path>
            <path style="stroke:#000000;stroke-width:126;stroke-linecap:round;"
              d="M 850,9259  A 4507 4507  0  0  1  4724  9448"></path>
            <polyline points="5385,4299 4062,8125"
              style="stroke:#000000;stroke-width:300; stroke-linejoin:miter; stroke-linecap:round;"></polyline>
            <polyline points="8598,3496 8598,9259"
              style="stroke:#000000;stroke-width:126; stroke-linejoin:miter; stroke-linecap:round;"></polyline>
            <polyline points="850,3496 850,9259"
              style="stroke:#000000;stroke-width:126; stroke-linejoin:miter; stroke-linecap:round;"></polyline>
            <polyline points="4960,9685 4488,9685"
              style="stroke:#000000;stroke-width:174; stroke-linejoin:miter; stroke-linecap:round;"></polyline>
            <polyline points="3070,4582 1889,6141 3070,7700"
              style="stroke:#000000;stroke-width:300; stroke-linejoin:miter; stroke-linecap:round;"></polyline>
            <polyline points="6418,4582 7600,6141 6418,7700"
              style="stroke:#000000;stroke-width:300; stroke-linejoin:miter; stroke-linecap:round;"></polyline>
            <polyline points="8976,3590 8976,9732"
              style="stroke:#000000;stroke-width:174; stroke-linejoin:miter; stroke-linecap:round;"></polyline>
            <path style="stroke:#000000;stroke-width:174;stroke-linecap:butt;"
              d="M 4960,9685 A 4228 4228  0  0  1  8976  9732"></path>
          </g>
        </svg></div>
    </a><a class="runestone-link" href="https://runestone.academy" title="Runestone Academy"><img class="logo"
        src="https://runestone.academy/runestone/static/images/RAIcon_cropped.png"></a><a class="mathjax-link"
      href="https://www.mathjax.org" title="MathJax"><img class="logo"
        src="https://www.mathjax.org/badge/badge-square-2.png"></a>
  </div>
</div>
</body>

</html>